<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2025/26
Prof. Dr. Heiner Giefers

**Übung 1 <br>
Es soll ein einfaches Programm geschrieben werden, das die Funktionsweise des Befehls `grep` unter Unix simuliert. Der Benutzer soll aufgefordert werden, einen regulären Ausdruck einzugeben, mit dem dann die Anzahl der Zeilen in einer Datei (Parameter `filename`) gezählt wird, die mit dem regulären Ausdruck übereinstimmen:** 

*$ python grep.py<br>
Gib eine RegEx an: ^Author <br>
mbox.txt hat 1798 Zeilen die auf den Ausdruck "^Author" passen*

**Geben Sie die Anzahl der gezählten Zeilen zurück.**

```python
from unittest import TestCase
_asserts = TestCase()
assert_equal = _asserts.assertEqual
assert_dict_equal = _asserts.assertDictEqual
```

<!-- grade_id: cell-5589f54f815efd9c, points: 0, grade: false, solution: true -->
```python
import re
def grep_simulation(filename):
    regEx = input("Gib eine RegEx an:")
    with open(filename, 'r') as file:
        finds = 0
        for line in file:
            line = line.strip()
            if re.search(regEx, line):
                print(line)
                finds +=1
        return finds       
```

<!-- grade_id: cell-d51d5d317e82c222, points: 1, grade: true, solution: false -->
```python
"""grep_simulation Tests"""


def myinput(inputs):
    def foo(a):
        print(a)
        input.counter += 1
        if input.counter > len(inputs) :
            raise "Nicht genügend inputs"
        else:
            return inputs[input.counter-1]
    foo.counter = 0
    return foo


fname = "files/mbox-short.mbox"

input_old = input
input = myinput(["^Author"])
r1 = grep_simulation(fname)
input = myinput(["^X-"])
r2 = grep_simulation(fname)
input = input_old

assert 27 == r1, "Der Regex wurde nicht korrekt ausgewertet"
assert 216 == r2, "Der Regex wurde nicht korrekt ausgewertet"


```

Output:
```
Gib eine RegEx an:
Author: stephen.marquard@uct.ac.za
Author: louis@media.berkeley.edu
Author: zqian@umich.edu
Author: rjlowe@iupui.edu
Author: zqian@umich.edu
Author: rjlowe@iupui.edu
Author: cwen@iupui.edu
Author: cwen@iupui.edu
Author: gsilver@umich.edu
Author: gsilver@umich.edu
Author: zqian@umich.edu
Author: gsilver@umich.edu
Author: wagnermr@iupui.edu
Author: zqian@umich.edu
Author: antranig@caret.cam.ac.uk
Author: gopal.ramasammycook@gmail.com
Author: david.horwitz@uct.ac.za
Author: david.horwitz@uct.ac.za
Author: david.horwitz@uct.ac.za
Author: david.horwitz@uct.ac.za
Author: stephen.marquard@uct.ac.za
Author: louis@media.berkeley.edu
Author: louis@media.berkeley.edu
Author: ray@media.berkeley.edu
Author: cwen@iupui.edu
Author: cwen@iupui.edu
Author: cwen@iupui.edu
Gib eine RegEx an:
X-Sieve: CMU Sieve 2.3
X-Authentication-Warning: nakamura.uits.iupui.edu: apache set sender to stephen.marquard@uct.ac.za using -f
X-Content-Type-Outer-Envelope: text/plain; charset=UTF-8
X-Content
```

**Übung 2 <br>
Es soll ein Programm geschrieben werden, das nach Zeilen der folgenden Form sucht:**

`New Revision: 39772`

**Die Zahl soll aus jeder der Zeilen mit einem regulären Ausdruck und der Methode `findall()` extrahiert werden. Es soll dann der Durchschnitt der Zahlen berechnet und als Ganzzahl zurückgegeben werden.**


<!-- grade_id: cell-39c2a2a83e599bad, points: 0, grade: false, solution: true -->
```python
import re
def find_by_regex(filename):
    regEx = "New Revision: [0-9]{5}"
    with open(filename, 'r') as file:
        finds = list(map(int, re.findall("[0-9]{5}" , " ".join(re.findall(regEx, file.read().strip())))))
    return int(sum(finds)/len(finds))
```

<!-- grade_id: cell-1fc4221d62a6df92, points: 1, grade: true, solution: false -->
```python
"""find_by_regex tests"""

fname = "files/mbox-short.mbox"
res = find_by_regex(fname)
assert type(res) == int, "Die Funktion soll ein int-Objekt zurückliefern"
assert_equal(39756, res)
```